# Guía: Calidad de datos
## Introducción

¿Qué construirá?

En este tutorial realizará un análisis de la calidad de los datos utilizando los DataFrames de PySpark

¿Qué necesita?

Los siguientes requisitos se encuentran instalados en la máquina virtual asignada a cada estudiante, específicamente en el ambiente de anaconda llamado "base". Recuerde que tiene a su disposición el tutorial de conexión a máquinas virtuales en la semana 1 de Coursera
1. Python 3 con pip instalado
2. Jupyter Labs
3. Paquetes: Pyspark (3.2.1), pandas (1.2.1), numpy (1.20.0) y matplotlib (3.3.4)
Otros:
1. Controlador Connector J(ya se encuentra configurado)
2. Acceso a servidor remoto MySQL con base de datos relacional "WWImportersTransactional". Recuerde que tiene a su disposición el tutorial de conexión remoto a Mysql en la semana 1 de Coursera

## Análisis de calidad de datos

La calidad de los datos consiste en validar la idoneidad de los datos analizando varias dimensiones, entre las cuales resaltamos:

- Completitud: Datos que no existen o no se conocen.
- Unicidad (Duplicidad): Datos que son idénticos en diferentes registros o registros con valores idénticos en atributos en los que no debería ocurrir
- Consistencia: Esta medida se define por la validez o integridad del dato, en la definición de su estructura, al interior de una fila o entre diferentes filas de la misma fuente o de diferentes fuentes. El manejo de unidades en los datos y el significado del dato son generadores de inconsistencias.
- Validez: a nivel de formato y de sentido de los datos más alla de sus valores

Para asegurar una calidad mínima de los datos en estas dimensiones hay que entender su origen y posteriormente realizar un proceso de validación de cada una de las dimensiones de perfilamiento que llevará a un entendimiento de los datos disponibles.

A continuación se definen dos funciones, la primera cuenta los valores vacios por columna de dataframe, la segunda retorna un diccionario de columnas con cardinalidad superior al 50%

Te invito a revisar los recursos de calidad de datos que hemos visto en el curso, a continuacion ejemplos daremos algunos ejemplos por cada dimensión para el caso del tutorial



### Unicidad y validez

En la siguiente función retorne la cantidad de datos que únicos que existen en la fuente de datos 1

```python
def unicidad():
    # Inserte acá su código para testear la unicidad
    return 5
```

<!-- grade_id: cell-b67bc1f0acaec1ee, points: 10, grade: true, solution: false -->
```python
assert unicidad()==5, "Error en la prueba de unicidad"
```

```python
def validez():
    # Inserte acá su código para la validez de los datos
    return 5
```

<!-- grade_id: cell-12bc7066d6047c53, points: 0, grade: true, solution: false -->
```python
assert validez()==5, "Error en la prueba de unicidad"
```

```python
def contar_vacios():
    resultado =6;
    #código que cuenta los vacíos de la columna "Columna" de los datos
    return resultado
```

<!-- grade_id: cell-502317516b54686b, points: 0, grade: true, solution: false -->
```python
assert contar_vacios()==6, "Error en la prueba de Contar vacíos"
```

```python
def cardinalidad():
    resultado = 11
    #Código que determina la cardinalidad de la columna "Columna 1" de la fuente de datos
    return resultado
```

<!-- grade_id: cell-a21a22870e398bb7, points: 0, grade: true, solution: false -->
```python
assert cardinalidad()==11, "Error en la prueba de cardinalidad"
```

<!-- grade_id: cell-752fa23e17b2546a, points: 0, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert cardinalidad()==11, "Error en la prueba de cardinalidad"
### END HIDDEN TESTS
```

```python

```